Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Lab 2: Pandas Overview

**This lab is due Thursday, 9/15/2020 at 11:59PM.**

[Pandas](https://pandas.pydata.org/) is one of the most widely used Python libraries in data science. In this lab, you will learn commonly used data wrangling operations and tools in Pandas. We aim to give you familiarity with:

* Creating dataframes
* Slicing data frames (i.e., selecting rows and columns)
* Filtering data (using boolean arrays)
* Data Aggregation and Grouping in dataframes

In this lab, you are going to use several pandas methods like `drop()`, `loc()`, `groupby()`. Remember that you can press `shift+tab` on any method to see the documentation for that method.

## Setup

```python
# Import the following packages. Note the shorthand for pandas.
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
%matplotlib inline
```

## Part 1: Creating DataFrames & Basic Manipulations

A [dataframe](https://pandas.pydata.org/pandas-docs/stable/user_guide/dsintro.html#dataframe) is a two-dimensional labeled data structure with columns holding data of potentially different types.

**Method 1: ** You can create a data frame by specifying the columns and values as shown below.

Notice the syntax: you're passing a dictionary into the DataFrame.  The keys become the column names (e.g. `'name'`), and the values are lists (`['Peter',....`)

```python
animals = pd.DataFrame(
    data={'name': ['Peter', 'Nutkin', 'Hunca Munca', 'Jemima'],
          'species': ['rabbit', 'squirrel', 'mouse', 'duck']
          })
animals
```

Output:
```
name   species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

**Method 2: ** You can also define a dataframe by specifying the rows like below.

Here, you're passing in tuples for each row of data (e.g. `("Peter", "rabbit")`) and specifying the column names separately.

```python
animals2 = pd.DataFrame(
    [("Peter", "rabbit"), ("Nutkin", "squirrel"), ("Hunca Munca", "mouse"),
     ("Jemima", "duck")], 
    columns = ["name", "species"])
animals2
```

Output:
```
name   species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

**Other methods**: Usually you won't be creating data frames in such a manual way.  You'll often be loading dataframes in from other file types -- for example, comma separated (csv) files.  More on that later.

You can obtain the dimensions of a dataframe by using the shape attribute, `dataframe.shape`

```python
(num_rows, num_columns) = animals.shape
num_rows, num_columns
```

Output:
```
(4, 2)
```

### Question 1

You can add a column using the syntax `dataframe['new column name'] = [data]`. Add a column called `favorite food` to the `animals` table which contains the strings 'nut', 'carrot', 'corn', and 'cheese'. Use your best guess as to which animal prefers which food. 

(note you'll need to comment out or delete the `NotImplementedError()`)

<!-- grade_id: q1_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
animals['favorite food'] = ['carrot','nut','cheese','corn']
# raise NotImplementedError()
```

```python
animals
```

Output:
```
name   species favorite food
0        Peter    rabbit        carrot
1       Nutkin  squirrel           nut
2  Hunca Munca     mouse        cheese
3       Jemima      duck          corn
```

### Question 2

Use the `.drop()` method to [drop](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.drop.html) the `favorite food` column you created, and save the new dataframe (without the `favorite food` column) to `animals_original`. Some notes:

* You'll need to look up `drop` to figure out the right syntax.
* Make sure to use the `axis` parameter correctly

<!-- grade_id: q2_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
animals_original = animals.drop(['favorite food'], axis=1)
#raise NotImplementedError()
```

```python
animals_original
```

Output:
```
name   species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert animals_original.shape[1] == 2
```

### Question 3

Use the `.rename()` method to [rename](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.rename.html) the columns of `animals_original` so they begin with a capital letter. Set the `inplace` parameter correctly to change the `animals_original` dataframe. (hint: in Question 2, `drop` creates and returns a new dataframe instead of changing `animals` because `inplace` by default is `False`)

<!-- grade_id: q3_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
animals_original.rename(columns = {"name": "Name", "species":"Species"}, inplace = True)
# raise NotImplementedError()
```

```python
animals_original
```

Output:
```
Name   Species
0        Peter    rabbit
1       Nutkin  squirrel
2  Hunca Munca     mouse
3       Jemima      duck
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert animals_original.columns[1] == 'Species' # the column number might be different for you
```

*Background*: For the curious, the field values you just worked with were inspired by [Beatrix Potter's](https://en.wikipedia.org/wiki/Beatrix_Potter) characters.

## Part 2: CalEnviroScreen Data
Now that we have learned the basics, we'll use Pandas to wrangle a real-world dataset. Specifically, we will be working with the [California Communities Environmental Health Screening Tool (CalEnviroScreen)](https://oehha.ca.gov/calenviroscreen/report/calenviroscreen-30), which uses demographic and environmental information to identify communities that are susceptible to various types of pollution. The various fields in this dataset contribute to the CES score, which reflects a community's environmental conditions and its vulnerability to environmental pollutants.

Your lab02 folder contains an Excel file downloaded from [here](https://oehha.ca.gov/media/downloads/calenviroscreen/document/ces3results.xlsx).

Start by running the cell below, which creates an Excel file object in Pandas that we can then inspect. The cell below shows you the sheet names in the spreadsheet. Documentation on Pandas' Excel methods can be found [here](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_excel.html). 

```python
# run this cell to import the Excel file and see the names of the tabs
filename = 'ces3results.xlsx'
xl = pd.ExcelFile(filename)
print(xl.sheet_names) # display a list of the sheets in the spreadsheet
```

Output:
```
['CES 3.0 (2018 Update)', 'Data Dictionary', 'Missing&NAData', 'Demographic profile']
```

Run the cell below to load the first sheet of the Excel file and assign it to the variable `ces3`. 

```python
ces3 = xl.parse(xl.sheet_names[0]) # display the first sheet as Pandas dataframe
ces3.head()
```

Output:
```
Census Tract  Total Population California County    ZIP  \
0    6019001100              3174           Fresno   93706   
1    6071001600              6133    San Bernardino  91761   
2    6019000200              3167           Fresno   93706   
3    6077000801              6692       San Joaquin  95203   
4    6019001500              2206           Fresno   93725   

  Nearby City \n(to help approximate location only)   Longitude   Latitude  \
0                                            Fresno -119.781696  36.709695   
1                                           Ontario -117.618013  34.057780   
2                                            Fresno -119.805504  36.735491   
3                                          Stockton -121.314524  37.940517   
4                                            Fresno -119.717843  36.681600   

   CES 3.0 Score   CES 3.0 Percentile CES 3.0 \nPercentile Range  ...  \
0      94.090246           100.000000   95-100% (highest scores)  ...   
1      90.67783
```

Note that the dataframe contains 57 columns, but Pandas truncates the number we are able to see at once. We can show all columns using the `pd.set_option` method.

```python
pd.set_option('display.max_columns', None)
ces3.head()
```

Output:
```
Census Tract  Total Population California County    ZIP  \
0    6019001100              3174           Fresno   93706   
1    6071001600              6133    San Bernardino  91761   
2    6019000200              3167           Fresno   93706   
3    6077000801              6692       San Joaquin  95203   
4    6019001500              2206           Fresno   93725   

  Nearby City \n(to help approximate location only)   Longitude   Latitude  \
0                                            Fresno -119.781696  36.709695   
1                                           Ontario -117.618013  34.057780   
2                                            Fresno -119.805504  36.735491   
3                                          Stockton -121.314524  37.940517   
4                                            Fresno -119.717843  36.681600   

   CES 3.0 Score   CES 3.0 Percentile CES 3.0 \nPercentile Range  \
0      94.090246           100.000000   95-100% (highest scores)   
1      90.677839         
```

Notice that this dataset doesn't include the units in many of the column headings. Let's take a look at a different sheet to get more information about what we're looking at.

Run the following cell to load the data dictionary.

```python
dd = xl.parse('Data Dictionary', header = 6)
dd.head(10)
```

Output:
```
Variable Name  \
0                                       Census Tract   
1                                   Total Population   
2                                  California County   
3                                                ZIP   
4  Nearby City \n(to help approximate location only)   
5                                          Longitude   
6                                           Latitude   
7                                      CES 3.0 Score   
8                                 CES 3.0 Percentile   
9                           CES 3.0 Percentile Range   

                                         Description CalEnviroScreen Category  
0                   Census Tract ID from 2010 Census                      NaN  
1                  2010 population in census tracts                       NaN  
2  California county that the census tract falls ...                      NaN  
3  Postal ZIP Code that the census tract falls wi...                      NaN  
4  City or nearby city
```

### Question 4
The length of a dataframe is equivalent to its number of rows. Find the length of `ces3`. What does each row represent?

```python
# YOUR CODE HERE
len(ces3)
```

Output:
```
8035
```

*YOUR ANSWER HERE*

*Each row represents observations for a unique census tract.*

## Slicing Data Frames - selecting rows and columns


### Selection Using Label

**Column Selection** 
To select a column of a `DataFrame` by column label, the safest and fastest way is to use the `.loc` [method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.loc.html). General usage looks like `frame.loc[rowname,colname]`. (Reminder that the colon `:` means "everything").

- You can also slice across columns. For example, `ces3.loc[:, 'ZIP':]` would select all rows in the column `ZIP` and every column to the right.

- *Alternative:* While `.loc` is invaluable when writing production code, it may be a little too verbose for interactive use. One recommended alternative is the `[]` method, which takes on the form `frame['colname']`.

**Row Selection**
Similarly, if we want to select a row by its label, we can use the same `.loc` method. In this case, the "label" of each row refers to the index (i.e., primary key) of the dataframe.

```python
#Example:
ces3.loc[100:110, 'ZIP']
```

Output:
```
100    90003
101    90810
102    92113
103    90011
104    90501
105    91732
106    91352
107    95358
108    90059
109    95203
110    90058
Name: ZIP, dtype: int64
```

```python
#Example:  Notice the difference between these two methods
ces3.loc[100:110, ['ZIP']]
```

Output:
```
ZIP
100  90003
101  90810
102  92113
103  90011
104  90501
105  91732
106  91352
107  95358
108  90059
109  95203
110  90058
```

The `.loc` method actually uses the index (the bolded, leftmost series in the dataframe) rather than the row position to perform the selection. In the previous example, it's just a coincidence that the `.loc` syntax matches that of the array slicing syntax - the index and row position aren't always the same value. For example, you could set your index to a non-numeric code, like census tract or other unique ID, if that's how you want to identify your records.

Alternatively, we can use [`.iloc`](http://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.iloc.html) to slice the dataframe using row location and column position.

See the following example:

```python
#Example: We change the index from 0,1,2... to the Census Tract column
df = ces3.set_index("Census Tract") # Why might we want to use Census Tract instead of County or City?
df.head()
```

Output:
```
Total Population California County    ZIP  \
Census Tract                                              
6019001100                3174           Fresno   93706   
6071001600                6133    San Bernardino  91761   
6019000200                3167           Fresno   93706   
6077000801                6692       San Joaquin  95203   
6019001500                2206           Fresno   93725   

             Nearby City \n(to help approximate location only)   Longitude  \
Census Tract                                                                 
6019001100                                              Fresno -119.781696   
6071001600                                             Ontario -117.618013   
6019000200                                              Fresno -119.805504   
6077000801                                            Stockton -121.314524   
6019001500                                              Fresno -119.717843   

               Latitude  CES 3.0 Score   CES 3.0 Perc
```

We can now lookup rows by name directly:

```python
df.loc[[6037205120, 6019000200], :]
```

Output:
```
Total Population California County    ZIP  \
Census Tract                                              
6037205120                3618       Los Angeles  90023   
6019000200                3167           Fresno   93706   

             Nearby City \n(to help approximate location only)   Longitude  \
Census Tract                                                                 
6037205120                                         Los Angeles -118.211796   
6019000200                                              Fresno -119.805504   

               Latitude  CES 3.0 Score   CES 3.0 Percentile  \
Census Tract                                                  
6037205120    34.018755      78.043685            99.823433   
6019000200    36.735491      85.970036            99.974776   

             CES 3.0 \nPercentile Range SB 535 Disadvantaged Community  \
Census Tract                                                             
6037205120     95-100% (highest scores)                        
```

However, if we want to access rows by location we will need to use the integer loc (`iloc`) accessor:

```python
#Example: 
# df.loc[2:5,"Year"] # You can't do this
df.iloc[1:4,6:9]
```

Output:
```
CES 3.0 Score   CES 3.0 Percentile CES 3.0 \nPercentile Range
Census Tract                                                               
6071001600        90.677839            99.987388   95-100% (highest scores)
6019000200        85.970036            99.974776   95-100% (highest scores)
6077000801        82.491521            99.962164   95-100% (highest scores)
```

### Question 4

Selecting multiple columns is easy using `.loc`.  You just need to supply a list of column names.  Select the `California County`,`Diesel PM Pctl`, and `PM2.5 Pctl` columns **in that order** from the `ces3` table.

<!-- grade_id: q4_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
dsl_O3_PM = ces3.loc[:, ["California County", "Diesel PM Pctl","PM2.5 Pctl"]]
#raise NotImplementedError()
```

```python
dsl_O3_PM.head()
```

Output:
```
California County  Diesel PM Pctl  PM2.5 Pctl
0           Fresno        95.544493   97.218064
1    San Bernardino       92.121966   93.637725
2           Fresno        95.420037   97.218064
3       San Joaquin       73.515868   84.019461
4           Fresno        58.220286   97.218064
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
assert dsl_O3_PM.shape == (8035, 3)
assert dsl_O3_PM.columns[1] == "Diesel PM Pctl"
```

As you may have noticed above, the .loc() method is a way to re-order the columns within a dataframe.

## Filtering Data

### Filtering with boolean arrays

Filtering is the process of removing unwanted material.  In your quest for cleaner data, you will undoubtedly filter your data at some point: whether it be for clearing up cases with missing values, culling out fishy outliers, or analyzing subgroups of your data set.  Note that compound expressions have to be grouped with brackets. Example usage looks like `df[df[column name] < 5]]`.

For your reference, some commonly used comparison operators are given below.

Symbol | Usage      | Meaning 
------ | ---------- | -------------------------------------
==   | a == b   | Does a equal b?
<=   | a <= b   | Is a less than or equal to b?
>=   | a >= b   | Is a greater than or equal to b?
<    | a < b    | Is a less than b?
&#62;    | a &#62; b    | Is a greater than b?
~    | ~p       | Returns negation of p
&#124; | p &#124; q | p OR q
&    | p & q    | p AND q
^  | p ^ q | p XOR q (exclusive or)

In the following we construct the DataFrame containing only census tracts in Sacramento County.

```python
ces3_SC = ces3[ces3['California County'] == 'Sacramento '] # Note the space after "Sacramento." This kind of quirk can be remedied with some simple data cleaning techniques, which we'll discuss in future lessons.
ces3_SC.head(5)
```

Output:
```
Census Tract  Total Population California County    ZIP  \
75     6067005301              1823       Sacramento   95811   
213    6067000700              2806       Sacramento   95814   
409    6067005205              2109       Sacramento   95826   
443    6067002200              4004       Sacramento   95818   
538    6067004502              4795       Sacramento   95823   

    Nearby City \n(to help approximate location only)   Longitude   Latitude  \
75                                         Sacramento -121.482520  38.591604   
213                                        Sacramento -121.501660  38.581871   
409                                        Sacramento -121.395510  38.537469   
443                                        Sacramento -121.507156  38.561282   
538                                        Sacramento -121.462926  38.502552   

     CES 3.0 Score   CES 3.0 Percentile CES 3.0 \nPercentile Range  \
75       68.638370            99.054105   95-100% (highest scores)   
```

### Question 5
Select the census tracts in Alameda county whose CES 3.0 Percentile is 90 or higher.

(If you use condition `p` & condition `q` to filter the dataframe, make sure to use `df[(p) & (q)]`)

Hint: The column names, like the county names, are not "clean." Try using the `.columns` method to look up the **exact** column names. 

<!-- grade_id: q5a_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
AC_highCES = ces3[(ces3["California County"] == "Alameda ") & (ces3[' CES 3.0 Percentile']>=90)]
#raise NotImplementedError()
```

```python
AC_highCES
```

Output:
```
Census Tract  Total Population California County    ZIP  \
245    6001409000              3552          Alameda   94621   
287    6001409100              2255          Alameda   94603   
300    6001408800              5547          Alameda   94621   
517    6001409200              3152          Alameda   94603   
701    6001407300              2598          Alameda   94601   
769    6001409500              3122          Alameda   94621   

    Nearby City \n(to help approximate location only)   Longitude   Latitude  \
245                                           Oakland -122.221368  37.720011   
287                                           Oakland -122.183500  37.732326   
300                                           Oakland -122.196942  37.758804   
517                                           Oakland -122.177787  37.729085   
701                                           Oakland -122.210924  37.762179   
769                                           Oakland -122.183766  37.750446
```

<!-- grade_id: q5a_test, points: 1, grade: true, solution: false -->
```python
assert len(AC_highCES) == 6
assert AC_highCES[" CES 3.0 Percentile"].max() == 96.9100769327784
```

## Data Aggregration (Grouping Data Frames)

### Question 6
To count the number of instances of a value in a `Series`, we can use the `value_counts()` [method](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.value_counts.html) as `df["col_name"].value_counts()`. 

Count the number of different census tracts in each Californial county. (You may use the `ces3` DataFrame created above.) In other words, compute the number of rows in the table for each county.

<!-- grade_id: q6a_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
num_censusTracts = ces3["California County"].value_counts()

#raise NotImplementedError()
```

```python
num_censusTracts
```

Output:
```
Los Angeles        2343
San Diego           627
Orange              582
Riverside           453
Santa Clara         372
San Bernardino      369
Alameda             360
Sacramento          317
Contra Costa        207
Fresno              199
San Francisco       195
Ventura             173
San Mateo           157
Kern                151
San Joaquin         139
Sonoma               99
Solano               96
Stanislaus           94
Monterey             93
Santa Barbara        89
Placer               84
Tulare               78
Marin                55
San Luis Obispo      53
Santa Cruz           52
Butte                51
Merced               49
Shasta               48
El Dorado            42
Yolo                 41
Napa                 40
Imperial             31
Humboldt             30
Kings                27
Madera               23
Sutter               21
Nevada               20
Mendocino            20
Lake                 15
Yuba                 14
Siskiyou             14
Tehama          
```

<!-- grade_id: q6a_test, points: 1, grade: true, solution: false -->
```python
assert num_censusTracts["Alameda "] == 360
assert num_censusTracts.sum() == len(ces3)
```

### Question 7a

A more versatile way to aggregate data is to use the `.groupby()` [function](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.groupby.html). Find the sum of `Imp. Water Bodies` for each `California County` in the `ces3` table. Use the syntax `df.groupby("col_name").sum()`.

<!-- grade_id: q7a_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sum_impH2O = ces3.groupby("California County").sum()

# raise NotImplementedError()
```

```python
sum_impH2O.sort_values(by = "Imp. Water Bodies", ascending=False)['Imp. Water Bodies']
```

Output:
```
California County
Los Angeles        6976
San Diego          2458
Orange             1489
Alameda            1438
Ventura            1260
Sacramento         1170
San Joaquin        1056
San Francisco      1049
Contra Costa        983
Monterey            775
Santa Clara         744
Riverside           652
San Mateo           597
Imperial            497
Santa Barbara       478
Sonoma              450
Stanislaus          422
Santa Cruz          362
Marin               343
Solano              334
San Bernardino      276
San Luis Obispo     245
Merced              222
Yolo                204
Butte               174
Sutter              159
Placer              153
Napa                116
El Dorado           103
Humboldt            103
Shasta               96
Fresno               94
San Benito           86
Siskiyou             70
Colusa               68
Tulare               67
Mendocino            66
Tehama               49
Nevada               49
Yuba                 49
Lake                 4
```

<!-- grade_id: q7a_test, points: 1, grade: true, solution: false -->
```python
assert sum_impH2O.loc["Los Angeles", "Imp. Water Bodies"] == 6976
assert sum_impH2O.sort_values(by = "Imp. Water Bodies", ascending=False).index[3] == "Alameda "
```

### Question 7b
Take a look at the the Data Dictionary. What does the sum of the `Imp. Water Bodies` column represent?  In the process you'll read about "buffers".  What is a buffer?

```python
# SCRATCH WORK HERE
```

*YOUR ANSWER HERE*

*The sum of the number of pollutants across all impaired water bodies within buffered distances to populated blocks in the **county**.* 

*Buffers are areas around an object in space (often a point).  They are often used to aggregate data (as in `groupby`: by averaging, taking the max, counts, etc.) in the vicinity of the object.*

### Question 7c

What do the values in `ZIP` represent in the dataframe `sum_impH2O`? Why is the column `CES 3.0 Percentile Range` no longer present in the dataframe?

*YOUR ANSWER HERE*

*`ZIP` is the sum of all zip codes associated with the census tracts in a given county. Although this isn't particularly meaningful to us, it's how Python interprets the `.sum()` method of `groupby()` - it takes the sum of all numeric values in the dataframe, grouped by the specified column (`California County`).*

*`CES 3.0 Percentile Range` is no longer present in the dataframe because the column does not contain numeric values. Python isn't able to sum the column's values, so it drops it. We could keep this column by grouping by both `California County` and `CES 3.0 Percentile Range` in the groupby - that way, we would get a dataframe that has the total number of pollutants in impaired water bodies by both county and CES 3.0 category.*

### Question 7d

Find the mean of `Poverty` for each county for census tracts with population greater than or equal to 3,000 and with a `Pollution Burden Pctl` above 85.


<!-- grade_id: q7b_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
poverty_mean = ces3[(ces3["Total Population"] >= 3000) & (ces3["Pollution Burden Pctl"] > 85)].groupby("California County").mean()

# raise NotImplementedError()
```

```python
poverty_mean.sort_values(by = "Poverty", ascending=False)['Poverty']
```

Output:
```
California County
Kings             73.200000
Butte             72.500000
Yolo              64.750000
Fresno            61.354348
Kern              57.033333
Merced            49.787500
San Diego         49.353333
Tulare            49.080000
Stanislaus        48.320930
Los Angeles       47.291725
Santa Cruz        46.500000
Riverside         44.760000
San Joaquin       44.663158
Madera            44.650000
Sacramento        44.450000
San Bernardino    43.964384
Imperial          40.100000
Ventura           37.160000
Santa Clara       36.650000
Orange            36.213235
Alameda           31.150000
San Mateo         21.966667
San Francisco     20.200000
Name: Poverty, dtype: float64
```

<!-- grade_id: q7b_test, points: 1, grade: true, solution: false -->
```python
assert np.round(poverty_mean.loc["Alameda ", "Poverty"],2) == 31.15
assert len(poverty_mean) == 23
```

### Question 7e

What does your output to 7d represent?  Dig in to the data a little further and tell us what you notice about the `Poverty` field values in the counties/tract combinations that show up in your result to 7d, versus the `Poverty` field values for all tracts?

*YOUR ANSWER HERE*

*From the CES documentation, poverty is the "Percent of the population living below two times the federal poverty
level (5-year estimate, 2011-2015)."  If one digs in to the data, you can find that the poverty percentile for those tracts with high CES scores is higher than the average percentile across all tracts.*  

#### You are done! Remember to submit this lab on bCourses in both html and ipynb formats after clicking Kernel -> Restart & Run All.