After you've completed the lab, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "" # put your full name here
COLLABORATORS = "" # put the name(s) of anyone you worked with here
```

---

# Lab 1: Getting started

**This assignment was distributed the week of August 31st. There is no due date for this assignment, but you should try to complete it and ask any questions in office hours or on the discussion forum before Lab 2, which will be distributed on Tuesday, September 8th.**

This is the first lab of ER 131 and is intended to ensure that you have your Jupyter Hub and/or local Notebook environments configured correctly.

# Imports

These packages will be used over and over in the course.  In your own coding you don't *have* to import `numpy` as `np` (and so on), but it is what most people use, and using those abbreviations will improve how well your code ports over to others' environments.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import re
%matplotlib inline
```

# Exercises

## Part A

Write a function that returns a list of numbers, such that $ x_i=ln(i) $ for $1\leq i \leq n$. 

* Don't worry about the case where $n <= 0$.
* Round the outputs to the nearest three decimal places.
* Hint: Use the NumPy functions [numpy.round](https://numpy.org/doc/stable/reference/generated/numpy.round_.html) and [numpy.log](https://numpy.org/doc/stable/reference/generated/numpy.log.html). 
* The assert statements below will be very picky on where spaces show up, and this narrows the range of solutions available to print out the list

Additionally, write the function so that it raises the `NotImplementedError` *unless* you give the function something to return *before* the `NotImplementedError` statement appears. You can use the following syntax to raise an error:

```python 
raise NotImplementedError()
```
If you finish early, try to come up with more than one way to code this.

<!-- grade_id: cell-9388c587d0e639ff, points: 0, grade: false, solution: true -->
```python
def natlogs(n):
    """
    This function takes an integer n>0 as input.
    It takes the natural log of each integer from 1 to n and returns a list,
    such that the ith element of the returned list equals ln(i), rounded to 
    three decimal places.
    """
    
    # YOUR CODE HERE
    
    raise NotImplementedError()
```

Your function should print `[0.0, 0.693, 1.099, 1.386, 1.609, 1.792, 1.946, 2.079, 2.197, 2.303]` for $n=10$.  
Check that it does:

```python
print(natlogs(10))
```

Check that `natlogs` returns the correct output for several inputs.  

Note, the `assert` syntax below will throw an `AssertionError` if the boolean statements evaluate to `False` and do nothing if the statements are `True`. (You want your statements to be true.)

<!-- grade_id: squares-tests, points: 1, grade: true, solution: false -->
```python
assert natlogs(1) == [0.0]
assert natlogs(2) == [0.0, 0.693]
assert natlogs(10) == [0.0, 0.693, 1.099, 1.386, 1.609, 1.792, 1.946, 2.079, 2.197, 2.303]
```

## Part B
Write a function that "cleans" a string by:
1.  making all letters uppercase, and
2.  trimming leading and ending spaces.

The function should take in a "raw" string as input and return a "clean" string.

<!-- grade_id: cell-b1fa3ce3a6566fbc, points: 0, grade: false, solution: true -->
```python
def clean(string):
    """Return cleaned string"""
    
    # YOUR CODE HERE
        
    raise NotImplementedError()
```

Your function should clean the address " 12 Oak Street  " to "12 OAK STREET":

```python
clean(" 12 Oak Street ")
```

Check that the function works as expected for several inputs.

<!-- grade_id: test-is-palindrome, points: 1, grade: true, solution: false -->
```python
assert clean("ER 131-lab section 101") == "ER 131-LAB SECTION 101"
assert clean(" 09/01/2020 ") == "09/01/2020"
assert clean("$1500 Dollars ") == "$1500 DOLLARS"
```

### Bonus question
Modify your `clean` function to also replace any non-alphanumeric characters (not A-Z or 0-9) with a space (hint: you'll have to use [regular expressions](https://docs.python.org/3/library/re.html) for this part. Check out sections 8.2 and 8.3 of the [DS 100 textbook](https://www.textbook.ds100.org/intro) for more explanation and examples of how to use "regex").

```python
def clean_bonus(string):
    """Return cleaned string"""
    
    # OPTIONAL: YOUR CODE HERE
    
```

```python
# # Uncomment this cell to test clean_bonus, if you attempted it

# assert clean_bonus("ER 131-lab section 101") == "ER 131 LAB SECTION 101"
# assert clean_bonus(" 09/01/2020 ") == "09 01 2020"
# assert clean_bonus("$1500 DOLLARS ") == "1500 DOLLARS"
```

## Part C

Run the following code to see a pretty plot.

```python
from mpl_toolkits.mplot3d import axes3d
import matplotlib.pyplot as plt


# Create a plot
u = np.linspace(1.5*np.pi, -1.5*np.pi, 100)
[x,y] = np.meshgrid(u, u)
z = np.cos( np.sqrt(x.flatten()**2 + y.flatten()**2) ).reshape(x.shape)


fig = plt.figure()
ax = fig.add_subplot(111, projection='3d')
# Plot a basic wireframe.
ax.plot_wireframe(x, y, z, rstride=10, cstride=10)
ax.view_init(elev=30., azim=50)
plt.savefig("figure1.pdf")
```

In line 12, what does the number '111' in the `add_subplot` function represent? Under what circumstances might you enter "121" here instead?

*YOUR ANSWER HERE*

What is the purpose of the `rstride` and `cstride` parameters in the `plot_wireframe` function?

*YOUR ANSWER HERE*

In line 15, what happens when you change the `elev` parameter in the `view_init` function? How about `azim`?

*YOUR ANSWER HERE*

The last line of code saves the figure as a pdf file. Can you find this file? Where is it stored?

*YOUR ANSWER HERE*

## Part D

Submit your assignment on bCourses. Again, this lab will not be graded, but this exercise is intended to provide you with a "low stakes" way to familiarlize yourself with the bCourses interface you will use to submit assignments.

**Step 1:** Save your assignment as a Jupyter notebook and as an html file. Click **Kernel** --> **Restart & Clear Output**. Then, click **Cell** --> **Run All**. Then, go to the toolbar and click **File** -> **Download as** -> **.html**. Now, click **File** -> **Download as** -> **notebook (.ipynb)**. Use the following naming convention for your files: [LAST NAME]_Lab01_ER131.

**Step 2:** On bCourses, navigate to ENERES 131 > Assignments > Lab 1: Introduction. Click the box that says "Submit Assignment."

**Step 3:** Select "Browse..." and choose the lab files you have downloaded. You may need to do this more than once to upload both the html and the ipynb files.

**Step 4:** Select "Submit Assignment"

## Hooray, you're done!